Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 2
## Problem 4 (15 marks)

One of the most important constants in mathematics is Euler's number $e$.
It can be defined in various ways, including the infinite series
\begin{align}
e & = \sum_{n=0}^\infty \frac 1 {n!} \\&= \frac 1 {0!} + \frac 1 {1!} + \frac 1 {2!} + \frac 1 {3!} + \frac 1 {4!} + \dotsb \\&= 1 + 1 + \frac 1 2 + \frac 1 6 + \frac 1 {24} + \dotsb,
\end{align}
where the **factorial** $n!$ of a positive integer $n$ is given by $n! = 1 \cdot 2 \cdot 3 \dotsb n$, and $0! = 1$. (For example, $4! = 1 \cdot 2 \cdot 3 \cdot 4 = 24$.)

Write a function `approximate_e` which takes as input an integer $N\geqslant 1$ and returns the value of the finite sum
$$
\frac 1 {0!} + \frac 1 {1!} + \frac 1 {2!} + \dotsb + \frac 1 {N!}.
$$

For example, `approximate_e(0)` should return $1$, `approximate_e(1)` should return $1 + 1$ (also known as $2$), and `approximate_e(3)` should return $1 + 1 + \frac 1 2 + \frac 16$, and so on. These values should be returned as floating point numbers.

*Hint.* Define an auxiliary function `factorial` to compute factorials and use that within your function `approximate_e`.


<!-- grade_id: cell-95651e1da74d1379, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def approximate_e(N):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ceabb65ce517fcef, points: 3, grade: true, solution: false -->
```python
assert approximate_e(0) == 1
assert approximate_e(1) == 2
```

<!-- grade_id: cell-3be66d91e6583320, points: 3, grade: true, solution: false -->
```python
assert approximate_e(2) == 1 + 1 + 1/2
assert approximate_e(3) == 1 + 1 + 1/2 + 1/6
```

<!-- grade_id: cell-52cf3f5d80af0e19, points: 3, grade: true, solution: false -->
```python
assert approximate_e(4) == 1 + 1 + 1/2 + 1/6 + 1/24
assert approximate_e(5) == 1 + 1 + 1/2 + 1/6 + 1/24 + 1/120
```

<!-- grade_id: cell-741b871fe8ff6e06, points: 6, grade: true, solution: false -->
```python

```