Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 2
## Problem 1 (10 marks)

Write a function `concat` which takes as input a list of strings `L` and returns the list obtained by concatening the strings in `L` (in their natural order).

For example, a call `concat([ 'hi', ' ', 'there'])` should return the string `'hi there'`.
If `L` is an empty list, then the call `concat(L)` should return an empty string.

*Hint.* Several problems from this assignment can be solved by a suitable use of the **accumulator pattern**.

<!-- grade_id: cell-5210d62f3570b944, points: 0, grade: false, solution: true -->
```python
def concat(L):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-dca945837c8bb2e4, points: 2, grade: true, solution: false -->
```python
assert concat([ 'hi', ' ', 'there']) == 'hi there'
```

<!-- grade_id: cell-7b661ca499b8b94b, points: 2, grade: true, solution: false -->
```python
assert concat([]) == ''
```

<!-- grade_id: cell-229956f463dbb4c7, points: 6, grade: true, solution: false -->
```python

```