Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 2
## Problem 5 (11 marks)

A string is called a **palindrome** if it reads the same backward as forward. For example, the strings `'HANNAH'` and `'BOB'` are palindromes, but `'ALICE'` is not. For our purposes, capitalisation matters, so we do not regard `'Anna'` as a palindrome.

Write a function `is_palindrome` which takes as input a string `text` and returns `True` or `False`, according to whether `text` is or is not a palindrome.

<!-- grade_id: cell-a452139482fb5885, points: 0, grade: false, solution: true -->
```python
def is_palindrome(text):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3de858e106a4c5fd, points: 5, grade: true, solution: false -->
```python
assert is_palindrome('HANNAH')
assert is_palindrome('BOB')
```

<!-- grade_id: cell-e8cf2a7c2e33d373, points: 4, grade: true, solution: false -->
```python
assert not is_palindrome('ALICE')
assert not is_palindrome('Anna')
```

<!-- grade_id: cell-c614a4a60ebfdaa2, points: 2, grade: true, solution: false -->
```python

```