Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 2
## Problem 3 (12 marks)

Write a function `grade` which takes as input an integer `score` between `0` and `100` and returns a string according to the following table.

|Score|Grade|
|---|---|
|70-100|H1|
|60-69|H2.1|
|50-59|H2.2|
|40-49|Pass|
|0-39|Fail|

For example, `grade(82)` should return the string`'H1'` and `grade(42)` should return the string `'Pass'`.

<!-- grade_id: cell-e4fc98063b4edbb1, points: 0, grade: false, solution: true -->
```python
def grade(score):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c43ff63860a25846, points: 2, grade: true, solution: false -->
```python
from builtins import all, range
assert all(grade(i) in ['H1', 'H2.1', 'H2.2', 'Pass', 'Fail'] for i in range(101))
```

<!-- grade_id: cell-74679001fa4b786f, points: 1, grade: true, solution: false -->
```python
assert grade(82) == 'H1'
```

<!-- grade_id: cell-b79bab1a30fe1545, points: 2, grade: true, solution: false -->
```python
assert grade(65) == 'H2.1'
assert grade(51) == 'H2.2'
```

<!-- grade_id: cell-cbe4c5d99346b81b, points: 1, grade: true, solution: false -->
```python
assert grade(40) == 'Pass'
```

<!-- grade_id: cell-f651c7c064a8ad5b, points: 1, grade: true, solution: false -->
```python
assert grade(39) == 'Fail'
```

<!-- grade_id: cell-8173082d06fed053, points: 5, grade: true, solution: false -->
```python

```