Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 2
## Problem 2 (12 marks)

Let $x_1,\dotsc,x_n$ be a sequence of numbers, where $n\geqslant 1$. The [**(arithmetic) mean**](https://en.wikipedia.org/wiki/Arithmetic_mean) or **average** of this sequence is the quantity
$$\frac 1 n \sum_{i=1}^n x_i = \frac{x_1 + x_2 + \dotsb + x_n}{n}.$$
Arithmetic means are used frequently in mathematics, statistics, and numerous other fields such as economics.

Write a function `mean` which takes as its input parameter a non-empty list `numbers` of numbers and returns the mean of the numbers contained in `numbers`.

<!-- grade_id: cell-db4378470dcf2ee6, points: 0, grade: false, solution: true -->
```python
def mean(numbers):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-542dc24914663623, points: 3, grade: true, solution: false -->
```python
assert mean([5]) == 5
assert mean([1,4,5]) == 10/3
assert mean([2,2,2,2]) == 2
assert mean([1,2,3,4]) == 10/4
```

<!-- grade_id: cell-4795d6006d1e7460, points: 9, grade: true, solution: false -->
```python

```