Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 1
## Problem 5 (10 marks)

Suppose that beginning at 6am on some day, we have waited for precisely 1700 hours. Assuming that daylight saving time has been abolished, write a Python expression which computes the current hour of the day as an integer between 0 and 23.

<!-- grade_id: cell-ec50a7c0ac23a218, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-db49bb576ac9c288, points: 2, grade: true, solution: false -->
```python
_answer = _
assert _answer in range(24)
```

<!-- grade_id: cell-e621acf1bf0475ff, points: 8, grade: true, solution: false -->
```python

```