Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 1
## Problem 6 (15 marks)

### Background reading: temperature conversion
The following function `fahrenheit_to_celsius` converts degrees Fahrenheit to degrees Celsius:

```python
def fahrenheit_to_celsius(fahrenheit):
    celsius = 5/9 * (fahrenheit - 32)
    return celsius
```

If this function is called in the form `fahrenheit_to_celsius(x)`, the value of `x` is assigned to the **parameter** `fahrenheit`. The body of the function is then executed. The **return** statement in the third line finishes execution of the function. The function call evaluates to the value passed to the return statement (i.e. the value of `celsius` in our example).

<hr>

### First task
One inch (in) coincides with 2.54 centimetres (cm). Using the above function as a template, in the following cell, write a function `in_to_cm` which converts inches to centimetres.

<!-- grade_id: cell-639e1de02be651e9, points: 0, grade: false, solution: true -->
```python
def in_to_cm(inches):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ec4dd5e124416f39, points: 1, grade: true, solution: false -->
```python
assert type(in_to_cm(4)) == float
```

<!-- grade_id: cell-d8c16e575cf67d07, points: 3, grade: true, solution: false -->
```python
assert in_to_cm(1) == 2.54
```

<!-- grade_id: cell-b0bc7cf098c7a0da, points: 5, grade: true, solution: false -->
```python

```

<hr>

### Second task
In the following cell, write a function `cm_to_in` which converts centimetres to inches.

<!-- grade_id: cell-0542581c31ac52b0, points: 0, grade: false, solution: true -->
```python
def cm_to_in(centimetres):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1fa085e0decce885, points: 1, grade: true, solution: false -->
```python
assert cm_to_in(2.54) == 1
```

<!-- grade_id: cell-8c58f9579a34b9a4, points: 5, grade: true, solution: false -->
```python

```