Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 1
## Problem 3 (10 marks)

Write a Python expression which evaluates to the string consisting of 13 letters `a`, followed by 17 letters `b`, followed by 19 letters `c`.

<!-- grade_id: cell-bf8ad5bb43f6a764, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-07eba31ad87e4570, points: 2, grade: true, solution: false -->
```python
_answer = _
assert isinstance(_answer, str)
```

<!-- grade_id: cell-e7cd359913ea5c51, points: 8, grade: true, solution: false -->
```python
assert _answer[0] == 'a'
assert _answer[-1] == 'c'
```