Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 1
## Problem 4 (10 marks)

Write a Python expression which computes $$\left(1 - \frac 2 {3^2 + 4^3}\right)^2$$ as a floating point number.

<!-- grade_id: cell-d3b71d15491a4dae, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-00caa3a120d52e42, points: 2, grade: true, solution: false -->
```python
_answer = _
assert isinstance(_answer, float)
```

<!-- grade_id: cell-a0bb5835e14dc4ff, points: 8, grade: true, solution: false -->
```python

```