Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 1
## Problem 2 (10 marks)

Write a Python expression which computes the number of seconds in a week.

<!-- grade_id: cell-9dcb3f844f16917c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ac583942bdc79be, points: 2, grade: true, solution: false -->
```python
_answer = _
assert isinstance(_answer, int)
assert _answer > 0
```

<!-- grade_id: cell-b9a0caaebf6f8ea1, points: 8, grade: true, solution: false -->
```python

```