Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 1
## Problem 1 (5 marks)

Write down a Python expression that evaluates to the integer 7.

<!-- grade_id: cell-86ac0c7bec13dbca, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Note 1.** The cells following your solutions are used to perform *tests*. If you run a cell containing your solution and then run the corresponding test cell after that, you should not receive any error messages. Typically, the test cells that you see will only validate the *shape* of your answer without checking the correctness of your solution for you. For example, the following test cell merely checks if your answer above is a positive integer.

<!-- grade_id: cell-e04d7f7ebccd815c, points: 1, grade: true, solution: false -->
```python
_answer = _
assert isinstance(_answer, int)
assert _answer > 0
```

**Note 2.** In addition to the tests that you can see, further *hidden tests* will be performed to mark your answer. The following seemingly empty cell actually tests if your answer from above was correct.

<!-- grade_id: cell-2df1a2cc08080b55, points: 4, grade: true, solution: false -->
```python

```