Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 4
## Problem 4 (12 marks)

A positive quantity (e.g. an investment or a population) grows at a constant rate of `rate` percent every year. Write a function `doubling_time` which takes as input a positive floating point number `rate` and which returns the *least* non-negative integer `k` such that the initial quantity will have (at least) doubled after `k` years. For example `doubling_time(100)` should be `1` because at an annual growth rate of 100%, the initial size will have doubled after precisely one year.  On the other hand, `doubling_time(50)` should be `2` since the initial population will have grown to 150% of its initial size after one year, and then to 225% of its initial size after a total of two years. If `rate` is less than or equal to `0`, then your function should execute the statement
```python
    raise ValueError
```

*Note.*
If $x_0$ denotes the initial quantity, then by the **law of exponential growth**, after $k$ years, it will have grown to $x_0 \bigl(1+\frac{\mathrm{rate}}{100}\bigr)^k$. The initial quantity $x_0$ is irrelevant for this problem and you can e.g. take it to be `1`. (Any positive number will work.)

*Hints.* You can e.g. solve this problem by searching for the doubling time using a `while` loop. Alternatively, you could derive a mathematical formula for the doubling time that depends on `rate`, and then evaluate that formula.

<!-- grade_id: cell-195037e9bbedece8, points: 0, grade: false, solution: true -->
```python
def doubling_time(rate):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-666a3ea6a6b12a66, points: 4, grade: true, solution: false -->
```python
assert doubling_time(100) == 1
assert doubling_time(50) == 2
assert doubling_time(25) == 4
assert doubling_time(12.5) == 6
```

<!-- grade_id: cell-b3520ec0f60bf88c, points: 2, grade: true, solution: false -->
```python
try:
    doubling_time(0)
except ValueError:
    pass
else:
    raise AssertionError
```

<!-- grade_id: cell-cb536dd4f890302a, points: 6, grade: true, solution: false -->
```python

```