Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 4
## Problem 5 (12 marks)

A **prefix** of a string `S` is any initial slice `S[0:N]` for `0 <= N <= len(S)`. We say that a string `S` is a [**Dyck word** ](https://en.wikipedia.org/wiki/Dyck_language) if and only if the following conditions are satisfied:
   * The only characters in `S` are the opening and closing parentheses `'('` and `')'`, respectively.
   * `S` contains the same numbers of opening and closing parentheses.
   * Every prefix of `S` contains at least as many opening as closing parentheses.

Note that empty strings are Dyck words.

Write a function `is_dyck_word` that takes as input a string `S` and returns `True` or `False`, depending on whether `S` is a Dyck word or not.

*Note.* Dyck words describe the possible ways of putting syntactically reasonable parentheses around expressions. They appear in numerous enumeration problems.

<!-- grade_id: cell-e6d471fe5b7f2390, points: 0, grade: false, solution: true -->
```python
def is_dyck_word(S):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-12bc6d5478eaf9a3, points: 2, grade: true, solution: false -->
```python
assert is_dyck_word('')
assert not is_dyck_word(' ')
```

<!-- grade_id: cell-18f5b9f42675034e, points: 3, grade: true, solution: false -->
```python
assert is_dyck_word('(()())')
assert not is_dyck_word('())')
assert not is_dyck_word(')(')
assert not is_dyck_word('))')
```

<!-- grade_id: cell-f5f29bcb886ff32d, points: 7, grade: true, solution: false -->
```python

```