Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 4
## Problem 1 (12 marks)

Write a function `find_keys` which takes as input parameters a dictionary `D` and an arbitrary value `x` and returns a list consisting of all those keys `key` of `D` with `D[key] == x`. If there are no such keys, then your function should return an empty list. The lists returned by your function should never contain duplicate items.

<!-- grade_id: cell-b22d5f7d3ddfd6cb, points: 0, grade: false, solution: true -->
```python
def find_keys(D, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1f7e372b8a9705f4, points: 3, grade: true, solution: false -->
```python
_D = { 1:'a', 2:'b', 3:'c', 4:'a', 5:'b', 6:'c', 7:'c', 8:'b', 9:'a', 'ten': 'a'}
_res = find_keys(_D, 'a')
assert len(_res) == 4
assert set(_res) == set([1,4,9,'ten'])
```

<!-- grade_id: cell-725a2a6b07f5b9bd, points: 2, grade: true, solution: false -->
```python
assert find_keys(_D, 'd') == []
```

<!-- grade_id: cell-38b415fefd2d506d, points: 7, grade: true, solution: false -->
```python

```