Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 4
## Problem 2 (12 marks)

The [Lucas numbers](https://en.wikipedia.org/wiki/Lucas_number) are the sequence of numbers $L_0,L_1,\dotsc$ defined by the relations

$$\begin{align*}
L_0 & = 2 \\
L_1 & = 1 \\
L_k & = L_{k-1} + L_{k-2} \text{ for } k > 1.
\end{align*}
$$

These numbers are closely related to the [Golden ratio](https://en.wikipedia.org/wiki/Golden_ratio).

Write a Python function `lucas_up_to` that takes a positive integer `n` as its input parameter and returns
a list `lucas` of length `n+1` such that `lucas[i]` is $L_i$ for each integer `i` from `0` to `n`.

<!-- grade_id: cell-f625a1b8917d4190, points: 0, grade: false, solution: true -->
```python
def lucas_up_to(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-27799c30846e109f, points: 3, grade: true, solution: false -->
```python
assert lucas_up_to(2) == [2,1,3]
```

<!-- grade_id: cell-a8899494cdb3cb80, points: 3, grade: true, solution: false -->
```python
assert lucas_up_to(6) == [2,1,3,4,7,11,18]
```

<!-- grade_id: cell-deeef9fb4f2a4ffd, points: 6, grade: true, solution: false -->
```python

```