Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 4
## Problem 3 (12 marks)

Call a dictionary `D` **invertible** if its values are unique and immutable. In other words, `D` is invertible if different keys have different values associated with them and if all values are immutable.
Suppose that `D` is invertible. Then an **inverse** of `D` is a dictionary `E` such that `D[key] == value` if and only if `E[value] == key`.
Write a function `invert_dictionary` which takes as its input parameter a dictionary `D` whose values are immutable and returns an inverse of `D` if `D` is invertible and `None`, otherwise.

For example, if `D == { 'A': 1, 'Z': 26 }`, then `invert_dictionary(D)` should return `{ 1: 'A', 26: 'Z' }`.
If, on the other hand, `D == {'A': 1, 'Z': 1 }`, then `invert_dictionary(D)` should return `None`.

<!-- grade_id: cell-141c521af9111ffc, points: 0, grade: false, solution: true -->
```python
def invert_dictionary(D):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e0e620ab52c7e441, points: 3, grade: true, solution: false -->
```python
assert invert_dictionary({ 'A': 1, 'Z': 26 }) == { 1: 'A', 26: 'Z' }
```

<!-- grade_id: cell-ecf7a24166cda0db, points: 2, grade: true, solution: false -->
```python
assert invert_dictionary({'A': 1, 'Z': 1 }) is None
```

<!-- grade_id: cell-eb18761c8efeb300, points: 7, grade: true, solution: false -->
```python

```