Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 3
## Problem 5 (12 marks)

Given a positive integer $n$, its **Collatz sequence** $C_0,C_1,\dotsc$ is defined using $C_0 = n$ and the following rules  for $n \geqslant 0$:
  * If $C_n$ is even, then $C_{n+1} = \frac{C_n} 2$.
  * If $C_n$ is odd, then $C_{n+1} = 3C_n + 1$.
  * If $C_n = 1$, then we stop so that $C_0,\dotsc,C_n$ is the complete Collatz sequence of $n$.

For example, the Collatz sequence of 5 is
$$5,16,8,4,2,1$$
and the Collatz sequence of 3 is
$$3,10,5,16,8,4,2,1.$$

The [Collatz conjecture](https://en.wikipedia.org/wiki/Collatz_conjecture) from 1937 asserts that the Collatz sequence of *any* positive integer is finite. Despite many serious attempts at proving it, this conjecture remains open! 

Assuming the validity of the Collatz conjecture, write a function `collatz_sequence` which takes as input parameter a positive integer `n` and returns the Collatz sequence of `n` as a list of integers (!). For example, `collatz_sequence(5)` should return the list `[5, 16, 8, 4, 2, 1]`.

*Notes*.
  * Should the Collatz conjecture be false, your function is allowed to never terminate for certain inputs. (If you find an example of this, you'll be famous!)
  * Use integer division `//` to avoid floating point arithmetic. Your function should return lists of integers.
  * Two years ago, [Fields medalist](https://en.wikipedia.org/wiki/Fields_Medal) [Terence Tao](https://en.wikipedia.org/wiki/Terence_Tao) obtained a major breakthrough related to the Collatz conjecture. See [this](https://www.quantamagazine.org/mathematician-terence-tao-and-the-collatz-conjecture-20191211) popular science article in [Quanta Magazine](https://www.quantamagazine.org/). Here is a link to a preprint of Tao's work: [arXiv:1909.03562](https://arxiv.org/abs/1909.03562)

<!-- grade_id: cell-58ce3d9db4439332, points: 0, grade: false, solution: true -->
```python
def collatz_sequence(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a1da1e289fae90f4, points: 1, grade: true, solution: false -->
```python
from builtins import all, int, type
assert all(type(x) == int for x in collatz_sequence(10))
```

<!-- grade_id: cell-b978aecd6f02755b, points: 2, grade: true, solution: false -->
```python
assert collatz_sequence(3) == [3,10,5,16,8,4,2,1]
```

<!-- grade_id: cell-b26f9628dc99f8c3, points: 2, grade: true, solution: false -->
```python
assert collatz_sequence(4) == [4,2,1]
```

<!-- grade_id: cell-67b4467193ebefef, points: 2, grade: true, solution: false -->
```python
assert collatz_sequence(5) == [5,16,8,4,2,1]
```

<!-- grade_id: cell-7f19db5a7f1546b9, points: 5, grade: true, solution: false -->
```python

```