Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 3
## Problem 2 (12 marks)

Write a function `three_digits` that takes as input parameter a non-zero digit `d` (an integer between `1` and `9`) and
returns a list of integers consisting of all three-digit numbers such that
  * the first (= leftmost) digit is at least `d`,
  * the third (= rightmost) digit is at least as large as the first digit, and
  * the second (= middle) digit is at least as large as the maximum of the first and third digit.
    
The lists returned by `three_digits` should **never** contain duplicate numbers.

For example:
  * `three_digits(9)` should return `[999]`,
  * `three_digits(8)` should return `[888, 898, 899, 999]` (or some rearrangement of that list), and
  * `three_digits(7)` should return `[777, 787, 788, 797, 798, 799, 888, 898, 899, 999]` (or some rearrangement of that list).

<!-- grade_id: cell-491678ef6bc9ca18, points: 0, grade: false, solution: true -->
```python
def three_digits(d):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-06030473c4f7fde9, points: 2, grade: true, solution: false -->
```python
assert three_digits(9) == [999]
```

<!-- grade_id: cell-b1bb425835d520e5, points: 3, grade: true, solution: false -->
```python
from builtins import sorted
assert sorted(three_digits(8)) == [888, 898, 899, 999]
assert sorted(three_digits(7)) == sorted([777, 787, 797, 788, 798, 799, 888, 898, 899, 999])
```

<!-- grade_id: cell-2d1db6ae63cf1bd8, points: 7, grade: true, solution: false -->
```python

```