Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 3
## Problem 3 (12 marks)

Leap years in the [Gregorian calendar](https://en.wikipedia.org/wiki/Gregorian_calendar) are defined as follows:

> Every year that is exactly divisible by four is a leap year, except for years that are exactly divisible by 100, but these centurial years are leap years if they are exactly divisible by 400. For example, the years 1700, 1800, and 1900 are not leap years, but the year 2000 is.

Write a function `is_leap_year` which takes as input a year `year` and returns `True` or `False`, according to whether `year` is a leap year or not.

<!-- grade_id: cell-19ef8917bc055e3e, points: 0, grade: false, solution: true -->
```python
def is_leap_year(year):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-960f5dbf4150f854, points: 3, grade: true, solution: false -->
```python
assert not is_leap_year(1700)
assert not is_leap_year(1800)
assert not is_leap_year(1900)
assert is_leap_year(2000)
```

<!-- grade_id: cell-b857d26fabd78a30, points: 3, grade: true, solution: false -->
```python
assert is_leap_year(2016)
assert not is_leap_year(2017)
assert not is_leap_year(2018)
assert not is_leap_year(2019)
assert is_leap_year(2020)
```

<!-- grade_id: cell-d39591b7e9c1d156, points: 6, grade: true, solution: false -->
```python

```