Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 3
## Problem 1 (12 marks)

A company pays their employees &euro;35 per hour for up to forty hours a week, and &euro;38 for every hour beyond that. Write a Python function `pay` that takes the number `hours` (a non-negative integer) of hours worked during a week as its input parameter and returns the total salary (in euros) to be paid for that week.

<!-- grade_id: cell-32dac30104e7bfcd, points: 0, grade: false, solution: true -->
```python
def pay(hours):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73c4062bd9d6f627, points: 6, grade: true, solution: false -->
```python
assert pay(0) == 0
assert pay(20) == 700
assert pay(40) == 1400
```

<!-- grade_id: cell-80015f7bfb1f003f, points: 6, grade: true, solution: false -->
```python
assert pay(45) == 1590
```