Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 3
## Problem 4 (12 marks)

A positive integer $n$ is a prime if and only if (a) $n > 1$ and (b) no integer $q$ with $2 \leqslant q \leqslant n-1$ divides $n$.
Write a function `is_prime` that takes as input parameter a positive integer `n` and returns `True` or `False` according to whether `n` is or is not prime.

<!-- grade_id: cell-94868dce12752739, points: 0, grade: false, solution: true -->
```python
def is_prime(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7d5ef86b1a188dbd, points: 1, grade: true, solution: false -->
```python
assert not is_prime(1)
```

<!-- grade_id: cell-0050dab3e7f13bb4, points: 1, grade: true, solution: false -->
```python
assert is_prime(2)
```

<!-- grade_id: cell-1fc176b2f59bdcb8, points: 2, grade: true, solution: false -->
```python
assert is_prime(3) and is_prime(5) and is_prime(7) and is_prime(11) and is_prime(13)
```

<!-- grade_id: cell-90d30baabd4c04c0, points: 2, grade: true, solution: false -->
```python
assert not(is_prime(4) or is_prime(6) or is_prime(8) or is_prime(9) or is_prime(10) or is_prime(12))
```

<!-- grade_id: cell-8045a4f745940070, points: 6, grade: true, solution: false -->
```python

```