Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 5
## Problem 2 (12 marks)

One of the most important constants in mathematics and physics is $\pi$, the ratio of a circle's circumference to its diameter. In this problem, you will write a Python function that approximates $\pi$ based on 
the [Leibniz formula](https://en.wikipedia.org/wiki/Leibniz_formula_for_%CF%80),
$$
\frac{\pi} 4 = 1 - \frac 1 3 + \frac 1 5 - \frac 1 7 + \frac 1 9 \mp \dotsb.
$$
Equivalently,
$$
\pi = 4 - \frac 4 3 + \frac 4 5 - \frac 4 7 + \frac 4 9 \mp \dotsb
$$

Write a function `approximate_pi` which takes as input an integer $N\geqslant 1$ and returns the (finite!) sum
$$
\frac 4 1 - \frac 4 3 + \frac 4 5 \mp \dotsb \pm \frac{4}{2N-1}.
$$

For example, `approximate_pi(1)` should return $4$, `approximate_pi(2)` should return $4 - \frac 4 3$, `approximate_pi(3)` should return $4 - \frac 4 3 + \frac 4 5$, and so on.

*Hint.* A call `approximate_pi(N)` should compute a sum consisting of precisely $N$ summands whose signs alternate. To implement the sign changes, note that if $i$ is an integer, then $(-1)^i = 1$ if $i$ is even and $(-1)^i = -1$ if $i$ is odd.

<!-- grade_id: cell-95651e1da74d1379, points: 0, grade: false, solution: true -->
```python
def approximate_pi(N):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ceabb65ce517fcef, points: 1, grade: true, solution: false -->
```python
assert approximate_pi(1) == 4
assert approximate_pi(3) == 4 - 4/3 + 4/5

```

<!-- grade_id: cell-52cf3f5d80af0e19, points: 1, grade: true, solution: false -->
```python
assert approximate_pi(2) == 4 - 4/3
assert approximate_pi(4) == 4 - 4/3 + 4/5 - 4/7
```

<!-- grade_id: cell-8a17cf45c6f62d8f, points: 2, grade: true, solution: false -->
```python
from builtins import abs
assert abs(approximate_pi(5) - (4 - 4/3 + 4/5 - 4/7 + 4/9)) < 1e-10
assert abs(approximate_pi(7) - (4 - 4/3 + 4/5 - 4/7 + 4/9 - 4/11 + 4/13)) < 1e-10
```

<!-- grade_id: cell-b3ee294d73718b41, points: 2, grade: true, solution: false -->
```python
from builtins import abs
assert abs(approximate_pi(6) - (4 - 4/3 + 4/5 - 4/7 + 4/9 - 4/11)) < 1e-10
assert abs(approximate_pi(8) - (4 - 4/3 + 4/5 - 4/7 + 4/9 - 4/11 + 4/13 - 4/15)) < 1e-10
```

<!-- grade_id: cell-741b871fe8ff6e06, points: 6, grade: true, solution: false -->
```python

```