Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 5
## Problem 4 (12 marks)

A (circular, solid) disk in the plane is completely determined by its radius and its centre (given as $x$- and $y$-coordinates). Define a class `Disk` for representing such disks.
Instances of your class should have the following methods:
  * `__init__(self, r, x, y)` which assigns `r` to the instance variable `radius` and assigns `x` and `y` to
  the instance variables `centre_x` and `centre_y`, respectively. If `r < 0`, then the method should execute the statement
  ```python
  raise ValueError
  ```
  A radius of `0` is allowed. (What is the geometric meaning of this?)
    
  * `diameter(self)` which computes and returns the diameter of a disk.

  * `area(self)` which computes and returns the area of a disk.
  
  * `distance(self, other)` which computes and returns the distance between two disks.
  Given disks $D_1$ and $D_2$, where $D_i$ has centre $(x_i,y_i)$ and radius $r_i$, the distance between $D_1$ and $D_2$ is the maximum of $0$ and
  $$\sqrt{(x_1 - x_2)^2 + (y_1 - y_2)^2} - r_1 - r_2.$$
  * `__repr__(self)` which returns a representation of a disk as a string. This method is already provided below so you do not need to write it yourself.
  
*Note.* Use the constant `pi` and the function `sqrt`. These are already imported for you below.

<!-- grade_id: cell-4c15ef31e969f048, points: 0, grade: false, solution: true -->
```python
from math import pi, sqrt

class Disk:
    def __init__(self, r, x, y):
        # YOUR CODE HERE
        raise NotImplementedError()

    def __repr__(self):
        return "Disk({}, {}, {})".format(self.radius, self.centre_x, self.centre_y)
    
    def diameter(self):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def area(self):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def distance(self, other):
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-ad4223ff421358ab, points: 2, grade: true, solution: false -->
```python
disk = Disk(1, 2, 3)
assert disk.radius == 1
assert disk.centre_x == 2
assert disk.centre_y == 3
```

<!-- grade_id: cell-fa266a4077107fce, points: 1, grade: true, solution: false -->
```python
Disk(0, 0, 0)

try:
    Disk(-1, 0, 0)
except ValueError:
    pass
else:
    raise AssertionError
```

<!-- grade_id: cell-d1f5468be0b3c850, points: 3, grade: true, solution: false -->
```python
d1, d2, d3 = Disk(2, 0, 0), Disk(2, 1, 0), Disk(2, 3, 4)
assert d1.distance(d2) == 0.0
assert d1.distance(d3) == 1.0
assert d2.distance(d2) == 0.0
```

<!-- grade_id: cell-974997b9893d5703, points: 3, grade: true, solution: false -->
```python
d1, d2 = Disk(1, 0, 0), Disk(2, 0, 0)
assert d1.diameter() == 2
assert d2.diameter() == 4
```

<!-- grade_id: cell-97f92acf43766f27, points: 3, grade: true, solution: false -->
```python
d1, d2 = Disk(1, 0, 0), Disk(2, 0, 0)
assert d1.area() == pi
assert d2.area() == 4 * pi
```