Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 5
## Problem 5 (12 marks)

This problem is about the game [Tic-Tac-Toe](https://en.wikipedia.org/wiki/Tic-tac-toe) (= noughts and crosses = Xs and Os). We represent a $3\times 3$ board of the game by a list of lists `board` whose entries correspond to positions according to the following table:


||||
|---|----|----|
| `board[0][0]` | `board[0][1]` | `board[0][2]`|
| `board[1][0]` | `board[1][1]` | `board[1][2]`|
| `board[2][0]` | `board[2][1]` | `board[2][2]`|

Here, each entry `board[i][j]` is one of the characters `'X'`, `'O'` (that's a capital letter O, not the number zero!), or `' '` (space), where the latter corresponds to an empty cell.

Each such board determines an outcome according to the following rules:
* If the board contains a horizontal, vertical, or diagonal line of three Xs, but not of three Os, then X wins.
* If the board contains a horizontal, vertical, or diagonal line of three Os, but not of three Xs, then O wins.
* If the board contains at least one such line of Xs and one such line of Os, then the board corresponds to a draw.
* Otherwise, the board does not determine a winner.


**Task.**
Write a function `tictactoe` which takes as input a list of lists `board` as above and returns a string which indicates whether `board` corresponds to X winning, O winning, a draw, or neither player winning. The possible return values of `tictactoe` should be the following strings:
* `'X wins'`
* `'O wins'`
* `'draw'`
* `'no winner'`

Your function should handle all $3\times 3$ boards comprised of the characters `'X'`, `'O'`, and `' '`, not merely those that actually arise when the game is played. For example, `tictactoe` should identify X as the winner of the following configuration, even though this cannot actually arise based on the usual rules of the game (why is that?):

|| ||
|---|---|---|
| X ||X|
| O | X | X|
| ||X|

*Hints.* You could e.g. solve this problem using a long series of `if` statements. If you wish to type less, I suggest you explore if you can exploit symmetries of the problem.

<!-- grade_id: cell-e649c9bb445f03e2, points: 0, grade: false, solution: true -->
```python
def tictactoe(board):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f2b0382aa13477b4, points: 2, grade: true, solution: false -->
```python
X,O,B = 'X', 'O', ' '
board = [[O,X,X],
         [X,O,O],
         [X,O,O]]
assert tictactoe(board) == 'O wins'
```

<!-- grade_id: cell-b4557a5353629ecb, points: 2, grade: true, solution: false -->
```python
X,O,B = 'X', 'O', ' '
board = [[X,X,X],
         [X,O,O],
         [X,O,O]]
assert tictactoe(board) == 'X wins'
```

<!-- grade_id: cell-7f9efc9e9c028caa, points: 2, grade: true, solution: false -->
```python
X,O,B = 'X', 'O', ' '
board = [[X,O,X],
         [B,O,B],
         [B,O,O]]
assert tictactoe(board) == 'O wins'
```

<!-- grade_id: cell-46f5f4dba5ff593a, points: 2, grade: true, solution: false -->
```python
X,O,B = 'X', 'O', ' '
board = [[X,O,B],
         [X,O,B],
         [X,O,B]]
assert tictactoe(board) == 'draw'
```

<!-- grade_id: cell-b77b8c16e19cdea4, points: 2, grade: true, solution: false -->
```python
X,O,B = 'X', 'O', ' '
board = [[X,O,B],
         [O,X,X],
         [O,X,B]]
assert tictactoe(board) == 'no winner'
```

<!-- grade_id: cell-f415519d896a02ab, points: 2, grade: true, solution: false -->
```python

```