Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 5
## Problem 3 (12 marks)

Write a function `all_sums` which takes a non-negative integer `n` as its input and returns a list `L` with the following properties:
* The length of `L` is `n + 1`.
* For each integer `k` from `0` up to `n`, the value `L[k]` is a list which contains precisely those pairs `[a, b]` of non-negative integers which satisfy `a + b == k`. (The order of the elements of `L[k]` does not matter, but these lists should not contain any duplicates, nor any elements other than those indicated before.)

For example, if `L == all_sums(3)`, then `L` should have length `4` and the following conditions should be satisfied:
* `L[0]` should be the list `[ [0,0] ]`. Indeed, `[0,0]` is the only pair of non-negative integers with sum `0`.
* `L[1]` should be the list `[ [0,1], [1,0] ]` or its rearrangement `[ [1,0], [0,1] ]`. Indeed, `[0,1]` and `[1,0]` are precisely the two pairs of non-negative integers with sum `1`.
* `L[2]` should be the list `[ [0,2], [1,1], [2,0] ]` (or some rearrangement of this list).
* `L[3]` should be the list `[ [0,3], [1,2], [2,1], [3,0] ]` (or some rearrangement of this list).

<!-- grade_id: cell-32160ff7b6d4e677, points: 0, grade: false, solution: true -->
```python
def all_sums(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-de8f38f3781b2de9, points: 4, grade: true, solution: false -->
```python
from builtins import len, sorted
_L = all_sums(3)
assert len(_L) == 4
assert sorted(_L[3]) == [[0, 3], [1, 2], [2, 1], [3, 0]]
assert sorted(_L[2]) == [[0, 2], [1, 1], [2, 0]]
assert sorted(_L[1]) == [[0, 1], [1, 0]]
```

<!-- grade_id: cell-9d5477bb84ee7088, points: 8, grade: true, solution: false -->
```python

```