Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your student ID number below. Please also delete all lines that say `raise NotImplementedError()`. Finally, make sure that your solutions are not commented out.

```python
ID = ""
```

---

### CS102/CS103
Tobias Rossmann<br>
School of Mathematical and Statistical Sciences<br>
NUI Galway

# Assignment 5
## Problem 1 (12 marks)

Write a function `censor` which takes as input a string `text` and a list of strings `banned`.
Your function should return a string which is obtained from `text` by replacing each occurrence of any string from `banned` in `text` by the string consisting of the same numbers of asterisks (`*`).

For example:
* `censor('This is appaling and outrageous.', ['appaling', 'outrageous'])` should return `'This is ******** and **********.'`

* `censor('I strongly dislike this problem.', ['dislike', 'hate'])` should return `'I strongly ******* this problem.'`

*Note.* You may assume that no string in `banned` occurs as a substring of any other string in `banned`.

<!-- grade_id: cell-d95c5c9b42bc22fe, points: 0, grade: false, solution: true -->
```python
def censor(text, banned):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-178841fcf7cdf09b, points: 2, grade: true, solution: false -->
```python
assert censor('This is appaling and outrageous.', ['appaling', 'outrageous']) == 'This is ******** and **********.'
```

<!-- grade_id: cell-cacdf7a8b50bebe7, points: 2, grade: true, solution: false -->
```python
assert censor('I strongly dislike this problem.', ['dislike', 'hate']) == 'I strongly ******* this problem.'
```

<!-- grade_id: cell-da0d134ae5620596, points: 8, grade: true, solution: false -->
```python

```